> This is a self-correcting activity generated by [nbgrader](https://nbgrader.readthedocs.io). Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Classification metrics

## Environment setup

```python
import numpy as np
```

## Accuracy

### Question

Complete the following function to compute accuracy.

<!-- grade_id: cell-6a21120c3e66ee57, points: 0, grade: false, solution: true -->
```python
def accuracy_score(y_true, y_pred):
    # YOUR CODE HERE
```

<!-- grade_id: cell-29d49954b89dc997, points: 1, grade: true, solution: false -->
```python
y_true = np.array([2, 0, 2, 2, 0, 1])
y_pred = np.array([0, 1, 2, 2, 0, 2])

acc = accuracy_score(y_true, y_pred)
print(f'Accuracy: {acc * 100:.2f}%')

assert acc == 0.5
```

## TODO

- Add other metrics: precision, recall, f1-score
- Add confusion matrix

```python

```