<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/tensors/ImageManagement.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


# Challenge: Image Management

## Instructions

This is a self-correcting notebook generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.

## Package setup

```python
# Import needed packages
# You may add or remove packages should you need them
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
from sklearn.datasets import load_sample_images

# Display plots inline and change plot resolution to retina
%matplotlib inline
%config InlineBackend.figure_format = 'retina'
```

## Loading samples images

```python
# Load samples images
images = np.asarray(load_sample_images().images)
print('Number of images: ' + str(len(images)))
print('Shape of images tensor: ' + str(images.shape))

first_image = images[0]
# Display first image
plt.imshow(first_image)

# Print details about first image
print('Shape of each image: ' + str(first_image.shape))
print('RGB values of top left image pixel: ' + str(first_image[0][0]))
```

## Question

Store in a variable named `rgb_values` the RGB values of the bottom right pixel.

<!-- grade_id: cell-9c0e895933db24f4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-1b314f0fcd01c0c8, points: 1, grade: true, solution: false -->
```python
print(rgb_values)
assert rgb_values.shape == (3,)
assert np.array_equal([15,24,7], rgb_values)
```

## Question

Reshape the previous `images` tensor into a 2D tensor.

<!-- grade_id: cell-a8b32b066d7c4993, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-875df848cbfa7fa1, points: 1, grade: true, solution: false -->
```python
# Assert new tensor dimensions
# 819840 = 427x640x3 
assert images.shape == (2, 819840)
# Assert RGB values of first pixel in first image
assert np.array_equal([174, 201, 231], images[0,:3])
# Assert RGB values of last pixel in first image
assert np.array_equal([15,24,7], images[0,819837:])
```

```python

```