<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/tensors/OneHotEncoding.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


# Challenge: One-Hot Encoding

## Instructions

This is a self-correcting notebook generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.

## Package setup

```python
# Import needed packages
# You may add or remove packages should you need them
import numpy as np
from keras.utils import to_categorical
```

## Creating a non-binary tensor

```python
# Definition of a 1D tensor (vector)
x = np.array([3, 1, 0, 2, 1])
print(x)
```

## Question

Encode the previous tensor following a binary representation (`one-hot encoding`).

<!-- grade_id: q05, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: t05, points: 1, grade: true, solution: false -->
```python
print(x)
# Assert one-hot encoding
assert x.shape == (5, 4)
assert np.array_equal([0,0,0,1], x[0])
assert np.array_equal([0,1,0,0], x[1])
assert np.array_equal([1,0,0,0], x[2])
assert np.array_equal([0,0,1,0], x[3])
assert np.array_equal([0,1,0,0], x[4])
```