<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/tensors/TensorManagement.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


# Challenge: Tensor Management

## Instructions

This is a self-correcting notebook generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.

## Package setup

```python
# Import needed packages
import numpy as np
```

## Question

Create a 2D tensor (a matrix) with dimensions (3,4) containing integer values of your choice. Store this tensor in a variable named `x`.

<!-- grade_id: q01, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: t01, points: 1, grade: true, solution: false -->
```python
print(x)
# Assert dimensions
assert x.ndim == 2
assert x.shape == (3, 4)
# Assert data type 
assert issubclass(x.dtype.type, np.integer)
```

## Question

Update the shape of the previous tensor so that it has dimensions (6,2).

<!-- grade_id: q02, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: t02, points: 1, grade: true, solution: false -->
```python
print(x)
# Assert tensor dimensions
assert x.ndim == 2
assert x.shape == (6, 2)
```

## Question

Change the type of the previous tensor values to `float32`.

<!-- grade_id: q03, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: t03, points: 1, grade: true, solution: false -->
```python
print(x)
# Assert data type
assert issubclass(x.dtype.type, np.floating)
```

## Question

Normalize the previous tensor so that it has a mean of 0 and a standard deviation of 1 on its first axis.

<!-- grade_id: q04, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: t04, points: 1, grade: true, solution: false -->
```python
print(x)
x_mean = x.mean(axis=0)
x_std = x.std(axis=0)
# Assert mean
assert np.abs(np.max(x_mean)) < 10**-6
# Assert standard deviation
assert np.abs(np.max(x_std-1)) < 10**-6
```