<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/tools/TensorManagement.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Tensor Management

## Package setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 10
%config InlineBackend.figure_format = 'retina'
```

```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import load_sample_images
```

## Part 1: Tensor Basics

### Question

Create a 2D tensor (a matrix) with dimensions (3,4) containing integer values of your choice. Store this tensor in a variable named `x`.

<!-- grade_id: q01, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: t01, points: 1, grade: true, solution: false -->
```python
print(x)

# Assert dimensions
assert x.ndim == 2
assert x.shape == (3, 4)

# Assert data type 
assert issubclass(x.dtype.type, np.integer)
```

### Question

Update the shape of the previous tensor so that it has dimensions (6,2).

<!-- grade_id: q02, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: t02, points: 1, grade: true, solution: false -->
```python
print(x)

# Assert tensor dimensions
assert x.ndim == 2
assert x.shape == (6, 2)
```

### Question

Change the type of the previous tensor values to `float32`.

<!-- grade_id: q03, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: t03, points: 1, grade: true, solution: false -->
```python
print(x)

# Assert data type
assert issubclass(x.dtype.type, np.floating)
```

## Part 2: Image Management

```python
# Load samples images
images = np.asarray(load_sample_images().images)
print(f'Number of images: {len(images)}. Images tensor: {images.shape}')

first_image = images[0]
# Display first image
plt.imshow(first_image)

# Print details about first image
print(f'First image: {first_image.shape}')
```

### Question

Store in variables respectively named `rgb_values_topleft` and `rgb_values_bottomright` the RGB values of the top-left and bottom-right pixels of the first image.

<!-- grade_id: cell-a0e5a557e72d9ead, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-bf8a78bb27fc556d, points: 1, grade: true, solution: false -->
```python
print(f'Top-left pixel: {rgb_values_topleft}')
assert rgb_values_topleft.shape == (3,)

print(f'Bottom-right pixel: {rgb_values_bottomright}')
assert rgb_values_bottomright.shape == (3,)
```

### Question

Reshape the previous `images` tensor into a 2D tensor.

<!-- grade_id: cell-d8d679c4bd424d6f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-62d8a1119a9ec8d5, points: 1, grade: true, solution: false -->
```python
# Assert new tensor dimensions
assert images.shape == (2, 819840)

# Assert RGB values of top-left in first image
assert np.array_equal(rgb_values_topleft, images[0,:3])

# Assert RGB values of bottom-right pixel in first image
assert np.array_equal(rgb_values_bottomright, images[0,819837:])
```