> This is a self-correcting activity generated by [nbgrader](https://nbgrader.readthedocs.io). Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Heart disease

In this activity, you'll use a small [dataset](https://archive.ics.uci.edu/ml/datasets/heart+Disease) provided by the Cleveland Clinic Foundation for Heart Disease. 

Each row describes a patient, and each column describes an attribute. You will use this information to predict whether a patient has heart disease.

Below is a description of each column.

|  Column  |                           Description                          |  Feature Type  | Data Type |
|:--------:|:--------------------------------------------------------------:|:--------------:|:---------:|
| Age      | Age in years                                                   | Numerical      | integer   |
| Sex      | (1 = male; 0 = female)                                         | Categorical    | integer   |
| CP       | Chest pain type (0, 1, 2, 3, 4)                                | Categorical    | integer   |
| Trestbpd | Resting blood pressure (in mm Hg on admission to the hospital) | Numerical      | integer   |
| Chol     | Serum cholestoral in mg/dl                                     | Numerical      | integer   |
| FBS      | (fasting blood sugar > 120 mg/dl) (1 = true; 0 = false)        | Categorical    | integer   |
| RestECG  | Resting electrocardiographic results (0, 1, 2)                 | Categorical    | integer   |
| Thalach  | Maximum heart rate achieved                                    | Numerical      | integer   |
| Exang    | Exercise induced angina (1 = yes; 0 = no)                      | Categorical    | integer   |
| Oldpeak  | ST depression induced by exercise relative to rest             | Numerical      | float     |
| Slope    | The slope of the peak exercise ST segment                      | Numerical      | integer   |
| CA       | Number of major vessels (0-3) colored by flourosopy            | Numerical      | integer   |
| Thal     | 3 = normal; 6 = fixed defect; 7 = reversable defect            | Categorical    | string    |
| Target   | Diagnosis of heart disease (1 = true; 0 = false)               | Classification | integer   |

## Environment setup

```python
import platform

print(f"Python version: {platform.python_version()}")
assert platform.python_version_tuple() >= ("3", "6")

import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
```

```python
# Setup plots
%matplotlib inline
plt.rcParams["figure.figsize"] = 10, 8
%config InlineBackend.figure_format = "retina"
sns.set()
```

<!-- grade_id: cell-eb38bba5eed58d48, points: 0, grade: true, solution: true -->
```python
import sklearn

print(f"scikit-learn version: {sklearn.__version__}")

# You may add other imports here as needed
# YOUR CODE HERE
```

## Step 1: loading the data

### Question

Load the dataset into a pandas DataFrame named `df_heart`.

<!-- grade_id: cell-afd9ddd5c64e787a, points: 0, grade: false, solution: true -->
```python
csv_url = "https://raw.githubusercontent.com/bpesquet/mlhandbook/master/_datasets/heart.csv"

# YOUR CODE HERE
```

<!-- grade_id: cell-3ba7578fa759b644, points: 1, grade: true, solution: false -->
```python
print(f"df_heart: {df_heart.shape}")

assert df_heart.shape == (301, 14)
```

## Step 2: prepare the data

### Question

Use the following cells to discover data.

<!-- grade_id: cell-b97b6c7563f889a8, points: 0, grade: true, solution: true -->
```python
# Print info about the dataset

# YOUR CODE HERE
```

<!-- grade_id: cell-3f78cf6edc4b3dd9, points: 0, grade: true, solution: true -->
```python
# Print the first 10 data samples

# YOUR CODE HERE
```

<!-- grade_id: cell-ca2d684382789293, points: 0, grade: true, solution: true -->
```python
# # Print descriptive statistics for all numerical attributes

# YOUR CODE HERE
```

<!-- grade_id: cell-a593bae8a5387a22, points: 0, grade: true, solution: true -->
```python
# Print distribution of target values

# YOUR CODE HERE
```

### Question

Use the following cells to prepare data for training:
- Split data between training and test sets with a 20% ratio.
- Store inputs and labels in the `x_train` and `y_train` variables.
- Preprocess training input data as needed.

<!-- grade_id: cell-b30578ddb1339003, points: 0, grade: false, solution: true -->
```python
# Split dataset between training and test

# YOUR CODE HERE
```

<!-- grade_id: cell-44f08b492f1af23e, points: 1, grade: true, solution: false -->
```python
print(f"Training dataset: {df_train.shape}")
print(f"Test dataset: {df_test.shape}")

assert df_train.shape == (240, 14)
assert df_test.shape == (61, 14)
```

<!-- grade_id: cell-8fb3ecece3668106, points: 0, grade: false, solution: true -->
```python
# Split training dataset between inputs and target

# YOUR CODE HERE
```

<!-- grade_id: cell-bc34229a83d80b4f, points: 1, grade: true, solution: false -->
```python
print(f"Training data: {df_x_train.shape}")
print(f"Training labels: {y_train.shape}")

assert df_x_train.shape == (240, 13)
assert y_train.shape == (240,)
```

<!-- grade_id: cell-664c1ad18297e830, points: 0, grade: true, solution: true -->
```python
# Print numerical and categorical features

# YOUR CODE HERE

# Print all values for the "thal" feature

# YOUR CODE HERE
```

<!-- grade_id: cell-f0c436744b2794ce, points: 0, grade: false, solution: true -->
```python
# Use pipelines to preprocess data

# YOUR CODE HERE
```

<!-- grade_id: cell-a5e4a800ccbbc0fc, points: 1, grade: true, solution: false -->
```python
# Print preprocessed data shape and first sample
print(f"x_train: {x_train.shape}")
print(x_train[0])

assert x_train.shape == (240, 15)
```

## Step 3: train and evaluate a model

### Question

Use the following cells to:
- Train a SGD classifier on the training data.
- Evaluate its accuracy using K-fold cross-validation.
- Compute the precision, recall and f1-score metrics.
- Plot its confusion matrix and ROC curve.

<!-- grade_id: cell-a433e228f0f0dcba, points: 0, grade: true, solution: true -->
```python
# Fit a SGD classifier to the training set

# YOUR CODE HERE
```

<!-- grade_id: cell-d88fdfbd9a2e77c0, points: 0, grade: false, solution: true -->
```python
# Use cross-validation to evaluate accuracy, using 3 folds
# Store the result in the cv_acc variable

# YOUR CODE HERE
```

<!-- grade_id: cell-31b0c320a7a2101a, points: 1, grade: true, solution: false -->
```python
print(f"CV accuracy: {cv_acc}")

assert np.mean(cv_acc) > 0.70
```

```python
# Plot the confusion matrix for a model and a dataset
def plot_conf_mat(model, x, y):
    with sns.axes_style("white"):  # Temporary hide Seaborn grid lines
        display = plot_confusion_matrix(
            model, x, y, values_format="d", cmap=plt.cm.Blues
        )
```

<!-- grade_id: cell-1cccbb273d746ded, points: 0, grade: true, solution: true -->
```python
# Plot confusion matrix for the SGD classifier

# YOUR CODE HERE
```

<!-- grade_id: cell-2a58d2664309c9f1, points: 0, grade: true, solution: true -->
```python
# Compute precision, recall and f1-score for the SGD classifier

# YOUR CODE HERE
```

<!-- grade_id: cell-6cb699959f8e5cc8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

## Bonus

Train another classifier and plot ROC curves for both.

```python

```