<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/CNN_CIFAR10.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Classify Common Images

In this kata, you'll train models to associate images representing common objects with their class.

The [CIFAR10](https://www.cs.toronto.edu/~kriz/cifar.html) dataset consists of 60,000 32x32 colour images in 10 classes, with 6,000 images per class. The classes are completely mutually exclusive. There are 50,000 training images and 10,000 test images.

![CIFAR10 images](images/cifar10.png)

## Package setup

```python
# The mlkatas package contains various utility functions required by all katas
!pip install mlkatas
```

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
import mlkatas
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 10
%config InlineBackend.figure_format = 'retina'
sns.set()
```

```python
# Import ML packages (edit this list if needed)
import tensorflow as tf
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense, Conv2D, MaxPooling2D, Flatten
from tensorflow.keras.datasets import cifar10
from tensorflow.keras.utils import to_categorical

print(f'TensorFlow version: {tf.__version__}')
print(f'Keras version: {tf.keras.__version__}')
```

## Step 1: Loading the data

### Question

* Load the CIFAR10 dataset included with Keras.
* Display the first 30 test images.

<!-- grade_id: cell-d5437d107fff93b7, points: 0, grade: false, solution: true -->
```python
# Load the data

# YOUR CODE HERE
```

<!-- grade_id: cell-8b2897cece69a922, points: 0, grade: false, solution: true -->
```python
# Plot the first 30 images

# YOUR CODE HERE
```

## Step 2: Training a dense neural network

### Question

* Prepare data for training with a dense network. Set apart the first 10,000 images as validation data.
* Train a model on the data to obtain a validation accuracy > 45%. Store the training history in a variable named `history`.

<!-- grade_id: cell-74322701546e38db, points: 0, grade: false, solution: true -->
```python
# Prepare data for training

# YOUR CODE HERE
```

<!-- grade_id: cell-84d0d9c1295f5442, points: 0, grade: false, solution: true -->
```python
# Create and train a model

# YOUR CODE HERE
```

```python
# Plot training history
mlkatas.plot_loss_acc(history)
```

<!-- grade_id: cell-fd3b689ba867dd0e, points: 3, grade: true, solution: false -->
```python
# Retrieve final validation accuracy
val_acc = history.history['val_acc'][-1]

# Assert final accuracy
assert val_acc > 0.45
```

```python
# Evaluate model performance on test data
test_loss, test_acc = model.evaluate(x_test, y_test)

print(f'Test accuracy: {test_acc * 100:.2f}%')
```

## Step 3: Training a convnet

### Question

* Prepare data for training with a convolutional network. Set apart the first 10,000 images as validation data.
* Train a model on the data to obtain a validation accuracy around 70%. Store the training history in a variable named `history`.

<!-- grade_id: cell-bd0d6be134baf68a, points: 0, grade: false, solution: true -->
```python
# Prepare data for training

# YOUR CODE HERE
```

<!-- grade_id: cell-222a6c2a933cfd3c, points: 0, grade: false, solution: true -->
```python
# Create a convnet

# YOUR CODE HERE
```

<!-- grade_id: cell-281bae8d7c8d4c9c, points: 0, grade: false, solution: true -->
```python
# Train the convnet

# YOUR CODE HERE
```

```python
# Plot training history
mlkatas.plot_loss_acc(history)
```

<!-- grade_id: cell-96829ee1e5cbe2d8, points: 3, grade: true, solution: false -->
```python
# Retrieve final validation accuracy
val_acc = history.history['val_acc'][-1]

# Assert final validation accuracy
assert val_acc > 0.68
```

```python
# Evaluate model performance on test data
test_loss, test_acc = model.evaluate(x_test, y_test)

print(f'Test accuracy: {test_acc * 100:.2f}%')
```

```python

```