<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/DecisionTree_Tuning.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Tuning a Decision Tree

In this kata, you'll search for the best hyperparameters for a `DecisionTreeClassifier` on a planar dataset.

## Package setup

```python
# The mlkatas package contains various utility functions required by all katas
!pip install mlkatas
```

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import mlkatas
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 8
%config InlineBackend.figure_format = 'retina'
sns.set()
```

<!-- grade_id: cell-470f64a6acb7b286, points: 0, grade: false, solution: true -->
```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import make_moons

# YOUR CODE HERE
```

## Step 1: Preparing the data

```python
# Generate moon-shaped, noisy data
x, y = make_moons(n_samples=1000, noise=0.4)

mlkatas.plot_planar_data(x, y)
```

## Step 3: Training a classifier

### Question¶

Split data and labels into the `x_train`, `x_test`, `y_train` and `y_test` variable using a 25% ratio.

<!-- grade_id: cell-ddaa4233cfda35cf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-832215c5196f7f4c, points: 2, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
print(f'x_test: {x_test.shape}. y_test: {y_test.shape}')

assert x_train.shape == (750, 2)
assert y_train.shape == (750,)
assert x_test.shape == (250, 2)
assert y_test.shape == (250,)
```

### Question

Create a `DecisionTreeClassifier` instance and store it into the model variable. Fit this model on the training data.

<!-- grade_id: cell-6f2a6cef93415c85, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-dc94a13ecad281a6, points: 0, grade: true, solution: false -->
```python
# Compute accuracy on training and test sets
train_acc = model.score(x_train, y_train)
test_acc = model.score(x_test, y_test)

print(f'Training accuracy: {train_acc * 100:.2f}%')
print(f'Test accuracy: {test_acc * 100:.2f}%')

assert train_acc > .95
assert test_acc > .75
```

```python
# Plot the decision boundary
mlkatas.plot_decision_boundary(lambda x: model.predict(x), x_train, y_train)
```

## Step 4: Tuning the classifier

### Question

Search for the best values of `max_leaf_nodes` and `min_samples_split` through a `GridSearchCV` instance, stored in the `grid_search_cv` variable.

The final test accuracy should be around 85%.

<!-- grade_id: cell-375a6bb1083de705, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-f96ff3d2a3d909d8, points: 0, grade: true, solution: false -->
```python
# Search for the best parameters with the specified classifier on training data
grid_search_cv.fit(x_train, y_train)
```

<!-- grade_id: cell-eb38cb5d8bac64c8, points: 0, grade: true, solution: false -->
```python
# Display the best found classifier
grid_search_cv.best_estimator_
```

<!-- grade_id: cell-25dd378795f665f4, points: 2, grade: true, solution: false -->
```python
# Compute accuracy on training and test sets
train_acc = grid_search_cv.score(x_train, y_train)
test_acc = grid_search_cv.score(x_test, y_test)

print(f'Training accuracy: {train_acc * 100:.2f}%')
print(f'Test accuracy: {test_acc * 100:.2f}%')
```

```python
# Plot the decision boundary
mlkatas.plot_decision_boundary(lambda x: grid_search_cv.predict(x), x_train, y_train)
```

```python

```