<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/algorithms/DNN_Iris.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Iris Dataset

[Iris](https://archive.ics.uci.edu/ml/datasets/iris) is a well-known multiclass dataset. It contains 3 classes of flowers with 50 examples each. There are a total of 4 features for each flower. The goal here is to identify a flower's class.

![](images/Iris-versicolor-21_1.jpg)

## Package setup

```python
# The mlkatas package contains various utility functions required by all katas
!pip install mlkatas
```

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
import mlkatas
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 10
%config InlineBackend.figure_format = 'retina'
sns.set()
```

```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import load_iris
import tensorflow as tf
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense
from tensorflow.keras.optimizers import *
from tensorflow.keras.utils import to_categorical

print(f'TensorFlow version: {tf.__version__}')
print(f'Keras version: {tf.keras.__version__}')
```

## Step 1: Loading the data

```python
# Load the Iris dataset included with scikit-learn
dataset = load_iris()

# Put data in a pandas DataFrame
df_iris = pd.DataFrame(dataset.data, columns=dataset.feature_names)
# Add target and class to DataFrame
df_iris['target'] = dataset.target
df_iris['class'] = dataset.target_names[dataset.target]
# Show 10 random samples
df_iris.sample(n=10)
```

### Question

Store training input data in a variable named `x_train` and training targets in a variable named `y_train`.

<!-- grade_id: cell-68cb077281b12fe7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-c35ec2e0b638d9ec, points: 1, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
print(f'Labels: {y_train}')

assert x_train.shape == (150,4)
assert y_train.shape == (150,)
```

## Step 2: Preparing the data

### Question

Reshape `y_train` to one-hot encode the targets in a (150,3) matrix.

<!-- grade_id: cell-2ee80aeb35907e0c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Show a sample of encoded targets
df_iris_labels = pd.DataFrame(y_train)
df_iris_labels.sample(n=10)
```

<!-- grade_id: cell-be7316c44aaa3e7c, points: 1, grade: true, solution: false -->
```python
print(f'y_train: {y_train.shape}')

assert y_train.shape == (150,3)
assert np.array_equal([1,0,0], y_train[0])
assert np.array_equal([0,1,0], y_train[50])
assert np.array_equal([0,0,1], y_train[100])
```

## Step 3: Training a model

### Question

Train a model on the data to obtain a training accuracy > 93%. Store the training history in a variable named `history`.

Tip: for best results, use the Adam optimizer with a learning rate of 0.1.

```python
model.compile(Adam(lr=0.1), loss='categorical_crossentropy', metrics=['accuracy'])
```

<!-- grade_id: cell-fd837bbf22683090, points: 0, grade: false, solution: true -->
```python
# Train a multiclass classifier on the data

# YOUR CODE HERE
```

```python
# Plot training history
mlkatas.plot_loss_acc(history)
```

<!-- grade_id: cell-a93a4cc26df606c8, points: 2, grade: true, solution: false -->
```python
# Retrieve final accuracy
final_acc = history.history['acc'][-1]

# Assert final accuracy
assert final_acc > 0.93
```

```python

```