<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/DNN_FashionMNIST.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Fashion-MNIST Dataset

[Fashion-MNIST](https://github.com/zalandoresearch/fashion-mnist) is analogous to the MNIST handwritten digits dataset. It consists of 60,000 28x28 grayscale images (pixel values between 0 and 255) of 10 fashion categories, along with a test set of 10,000 images. The goal here is to associate a fashion image with its class.

![](images/fashion-mnist-sprite.png)

Each training and test example is assigned to one of the following label (fashion category):

| Label | Description |
| --- | --- |
| 0 | T-shirt/top |
| 1 | Trouser |
| 2 | Pullover |
| 3 | Dress |
| 4 | Coat |
| 5 | Sandal |
| 6 | Shirt |
| 7 | Sneaker |
| 8 | Bag |
| 9 | Ankle boot |

## Package setup

```python
# The mlkatas package contains various utility functions required by all katas
!pip install mlkatas
```

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
import mlkatas
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 10
%config InlineBackend.figure_format = 'retina'
sns.set()
```

```python
# Import ML packages (edit this list if needed)
import tensorflow as tf
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense
from tensorflow.keras.datasets import fashion_mnist # Added in Keras 2.0.9
from tensorflow.keras.utils import to_categorical

print(f'TensorFlow version: {tf.__version__}')
print(f'Keras version: {tf.keras.__version__}')
```

## Step 1: loading the data

### Question

Load the Fashion-MNIST dataset included with Keras. Use variables named `train_images`, `train_labels`, `test_images` and `test_labels` to store the data.

<!-- grade_id: cell-7adc7072e5bfe5de, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-3ded7b39722f1cce, points: 1, grade: true, solution: false -->
```python
print(f'Training images: {train_images.shape}. Training labels: {train_labels.shape}')
print(f'Training labels: {train_labels}')
assert train_images.shape == (60000, 28, 28)
assert train_labels.shape == (60000, )
assert test_images.shape == (10000, 28, 28)
assert test_labels.shape == (10000, )
```

### Question

Display the first 10 test images.

<!-- grade_id: cell-1e4d9007e57da001, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

## Step 2: preparing the data

### Question

Create variables `x_train` and `x_test` by reshaping training and test images into 2D tensors with all values between 0 and 1.

<!-- grade_id: cell-215eb4b002815876, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-ea6880f444239bec, points: 1, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. x_test: {x_test.shape}')
# Assert tensors shapes
assert x_train.shape == (60000, 784)
assert x_test.shape == (10000, 784)
# Assert tensors values
assert (np.amin(x_train, axis=0) >= 0).all() and (np.amax(x_train, axis=0) <= 1).all()
assert (np.amin(x_test, axis=0) >= 0).all() and (np.amax(x_test, axis=0) <= 1).all()
```

### Question

One-hot encode training and test targets into variables `y_train` and `y_test`.

<!-- grade_id: cell-60c6a493ea4e216a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Show a sample of encoded targets
df_y_train = pd.DataFrame(y_train)
df_y_train.sample(n=10)
```

<!-- grade_id: cell-9234d933df4cf805, points: 1, grade: true, solution: false -->
```python
print(f'y_train: {y_train.shape}. y_test: {y_test.shape}')

# Assert one-hot encoding of training and test targets
assert y_train.shape == (60000, 10)
assert y_test.shape == (10000, 10)

# Assert samples values
assert np.array_equal([0,0,0,0,0,0,0,0,0,1], y_train[0])
assert np.array_equal([1,0,0,0,0,0,0,0,0,0], y_train[1])
assert np.array_equal([0,0,0,0,0,0,0,0,0,1], y_test[0])
assert np.array_equal([0,0,0,0,0,1,0,0,0,0], y_test[9999])
```

## Step 3: training a model 

### Question

Train a model on the data to obtain a test accuracy > 84%. Store the training history in a variable named `history`.

<!-- grade_id: cell-7f12afb21e4d1c88, points: 0, grade: false, solution: true -->
```python
# Create and train a model

# YOUR CODE HERE
```

```python
# Plot training history
mlkatas.plot_loss_acc(history)
```

<!-- grade_id: cell-54e0582707f93d74, points: 2, grade: true, solution: false -->
```python
# Evaluate the model on test data
test_loss, test_acc = model.evaluate(x_test, y_test, verbose=0)

print(f'Test accuracy: {test_acc * 100}%')

assert test_acc > 0.84
```

### Question

Show the first 10 test images with their predicted category ("dress", "shirt", etc).

<!-- grade_id: cell-4f6c6156250bb4ff, points: 0, grade: false, solution: true -->
```python
# Item descriptions, indexed by label value
descriptions = ['t-shirt', 'trouser', 'pullover', 'dress',
                'coat', 'sandal', 'shirt', 'sneaker', 'bag', 'ankle boot']

# YOUR CODE HERE
```

```python

```