<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/Ensemble_MNIST.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Ensemble Learning on MNIST

In this kata, you'll try several classifiers on the MNIST handwritten digits datasets, either separately and into an ensemble.

## Package setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 8
%config InlineBackend.figure_format = 'retina'
sns.set()
```

<!-- grade_id: cell-5c9c6ed91265b8ff, points: 0, grade: false, solution: true -->
```python
# Import ML packages (edit this list if needed)
from tensorflow.keras.datasets import mnist

# YOUR CODE HERE
```

## Step 1: Loading and preparing the data

### Question

Load the MNIST dataset included with Keras. Do the necessary preprocessing and split the data into training, validation and test sets.

<!-- grade_id: cell-fb6a9c746b43c1af, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-60dcf8104737b392, points: 2, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. x_val: {x_val.shape}. x_test: {x_test.shape}')
print(f'y_train: {y_train.shape}. y_val: {y_val.shape}. y_test: {y_test.shape}')

assert x_train.shape == (50000, 784)
assert x_val.shape == (10000, 784)
assert x_test.shape == (10000, 784)
assert y_train.shape == (50000,)
assert y_val.shape == (10000,)
assert y_test.shape == (10000,)
```

## Step 2: Train several models

### Question

Create and train various models, such as a linear classifier, a multilayer perceptron, a random forest...

<!-- grade_id: cell-9cf3e82d97c98c75, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Display the default score for each model.

<!-- grade_id: cell-4f60774a06371248, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Create a `VotingClassifier` including all your models. Fit it on the train data.

<!-- grade_id: cell-ccc315abafd8246f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Show the `VotingClassifier` score and compare to each model's individual score.

<!-- grade_id: cell-560a77d9f9b57409, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Show the score for a soft voting classifier.

<!-- grade_id: cell-ff3fc8fba60b9acd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Compute the `VotingClassifier` score on the test data. Compare it to each model's individual score.

<!-- grade_id: cell-7a9a6c423c33ba0d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```