<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/algorithms/KNN_BreastCancer.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Breast Cancer Dataset

The [Breast Cancer][1] dataset is used for multivariate binary classification between benign and maligant tumors. There are 569 total samples with 30 features each. Features were computed from a digitized image of a fine needle aspirate of a breast mass. They describe characteristics of the cell nuclei present in the image.

![](images/breast-cancer-logo.jpg)

[1]: https://archive.ics.uci.edu/ml/datasets/Breast+Cancer+Wisconsin+(Diagnostic)

## Package setup

```python
# The mlkatas package contains various utility functions required by all katas
!pip install mlkatas
```

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
import mlkatas
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 10
%config InlineBackend.figure_format = 'retina'
sns.set()
```

```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import load_breast_cancer
from sklearn.linear_model import LogisticRegression
import tensorflow as tf
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense

print(f'TensorFlow version: {tf.__version__}')
print(f'Keras version: {tf.keras.__version__}')
```

## Step 1: Loading the data

### Question

* Load the Breast Cancer dataset included with scikit-learn. Store it in a variable named `dataset`. 
* Display 10 random samples with feature names, target and class.

<!-- grade_id: cell-438cd8e66cdd8522, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

## Step 2: Using a simple model

```python
# Use scikit-learn's builtin logistic regression classifier to obtain a pretty good accuracy
model = LogisticRegression()
model.fit(dataset.data, dataset.target)

accuracy = model.score(dataset.data, dataset.target)
print('Accuracy: {:.2f}%'.format(accuracy * 100))
```

## Step 3: Training a neural network model

### Question

Train a model on the data to obtain a training accuracy > 85%. Store the training history in a variable named `history`.

<!-- grade_id: cell-400da8c44b9cae45, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
mlkatas.plot_loss_acc(history)
```

<!-- grade_id: cell-c04f242e72fc4600, points: 1, grade: true, solution: false -->
```python
# Retrieve final accuracy
final_acc = history.history['acc'][-1]

# Assert final accuracy
assert final_acc > 0.85
```

## Step 4: Beating the simple model

### Question

Optimize training to beat the simple model by achieving a training accuracy > 96%.

<!-- grade_id: cell-b288a6495025292c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Plot training history
mlkatas.plot_loss_acc(history)
```

<!-- grade_id: cell-cbe1f1188753f1ed, points: 2, grade: true, solution: false -->
```python
# Retrieve final accuracy
final_acc = history.history['acc'][-1]

# Assert final accuracy
assert final_acc > 0.96
```

```python

```