<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/LinearRegression_Scratch.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Code Linear Regression From Scratch

In this kata, you'll code various linear regression algorithms and compare them to some of scikit-learn's linear regressors.

## Package setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 8
%config InlineBackend.figure_format = 'retina'
sns.set()
```

```python
# Import ML packages (edit this list if needed)
from sklearn.linear_model import LinearRegression, SGDRegressor
```

## Generate planar data

```python
true_intercept = 4
true_slope = 3

# Generate linear-looking data with noise
X = 2 * np.random.rand(100, 1)
y = 4 + 3 * X + np.random.randn(100, 1)

# add x0 = 1 to each sample
X = np.c_[np.ones((100, 1)), X]

print(f'X: {X.shape}. y: {y.shape}')
```

```python
# Plot data (only x1)
plt.plot(X[:,1], y, "b.")
plt.xlabel("$x_1$", fontsize=18)
plt.ylabel("$y$", rotation=0, fontsize=18)
plt.axis([0, 2, 0, 15])
```

## Analytical Approach

### Question

Using the normal equation, compute the theta vector that minimizes loss into the `theta_best` variable.

<!-- grade_id: cell-8b0138760e3b7131, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-3b660b712c1c825c, points: 1, grade: true, solution: false -->
```python
normal_intercept = theta_best[0][0]
normal_slope = theta_best[1][0]
print(f'Intercept: {normal_intercept}. Slope: {normal_slope}')

assert np.abs(true_intercept - normal_intercept) < 0.5
assert np.abs(true_slope - normal_slope) < 0.5
```

```python
# Generate 2 test samples
X_test = np.array([[0], [2]])
# add x0 = 1 to each sample
X_test = np.c_[np.ones((2, 1)), X_test]
```

### Question

Use the model to predict values for the test samples. Store the result in variable `y_pred`.

<!-- grade_id: cell-6264901b299337ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-c937e69e2311ea21, points: 0, grade: true, solution: false -->
```python
print(f'y_pred: {y_pred}')
```

```python
plt.plot(X_test[:,1], y_pred, "r-", linewidth=2, label="Predictions")
plt.plot(X[:,1], y, "b.")
plt.xlabel("$x_1$", fontsize=18)
plt.ylabel("$y$", rotation=0, fontsize=18)
plt.legend(loc="upper left", fontsize=14)
plt.axis([0, 2, 0, 15])
```

### Question

Create a scikit-learn `LinearRegression` into the variable `model`. Use it to fit the data.

Afterwards, compare its paramaters to those of your regressor.

<!-- grade_id: cell-d9153fb070dfa6ab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-486ba599c756e970, points: 1, grade: true, solution: false -->
```python
sklearn_intercept = model.intercept_[0]
sklearn_slope = model.coef_[0][1]
print(f'Intercept: {sklearn_intercept}. Slope: {sklearn_slope}')

assert np.abs(sklearn_intercept - normal_intercept) < 0.5
assert np.abs(sklearn_slope - normal_slope) < 0.5
```

### Question

Use the scikit-learn model to predict and print values for the test samples.

<!-- grade_id: cell-b44a7432db69ee83, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

## Iterative Approach: Batch Gradient Descent

### Question

Complte the following function to implement Mean Squared Error loss.

<!-- grade_id: cell-b2b5eba005be50eb, points: 0, grade: false, solution: true -->
```python
def mse(y_true, y_pred):
    # YOUR CODE HERE
```

### Question

Complte the following code to implement batch gradient descent on the dataset, using the given hyperparameters.

Print the loss value every 10 iterations.

<!-- grade_id: cell-d61bd7e54acc65dc, points: 0, grade: false, solution: true -->
```python
eta = 0.1 # learning rate
n_iterations = 100
m = 100

theta_batch = np.random.randn(2,1) # random initialization
print(f'Initial cost: {mse(y, X.dot(theta_batch))}')

# YOUR CODE HERE
```

<!-- grade_id: cell-4f733fd7932337a0, points: 1, grade: true, solution: false -->
```python
batch_intercept = theta_batch[0][0]
batch_slope = theta_batch[1][0]
print(f'Intercept: {batch_intercept}. Slope: {batch_slope}')

assert np.abs(true_intercept - batch_intercept) < 0.5
assert np.abs(true_slope - batch_slope) < 0.5
```

## Iterative Approach: Stochastic Gradient Descent

### Question

Complte the following code to implement stochastic gradient descent on the dataset, using the given hyperparameters.

Print the loss value every 25 iterations.

<!-- grade_id: cell-1cb6981295853e52, points: 0, grade: false, solution: true -->
```python
eta = 0.01 # learning rate
n_iterations = 500
m = 100

theta_stochastic = np.random.randn(2,1) # random initialization
print(f'Initial cost: {mse(y, X.dot(theta_stochastic))}')

# YOUR CODE HERE
```

<!-- grade_id: cell-06537d7ad534e193, points: 1, grade: true, solution: false -->
```python
stochastic_intercept = theta_batch[0][0]
stochastic_slope = theta_batch[1][0]
print(f'Intercept: {stochastic_intercept}. Slope: {stochastic_slope}')

assert np.abs(true_intercept - stochastic_intercept) < 0.5
assert np.abs(true_slope - stochastic_slope) < 0.5
```

### Question

Create a scikit-learn `SGDRegressor` into the variable `model`. Use it to fit the data with the same hyperparameters as before.

Afterwards, compare its paramaters to those of your regressor.

<!-- grade_id: cell-4270a529be207e00, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-de18eaea25268f39, points: 1, grade: true, solution: false -->
```python
sklearn_intercept = model.intercept_[0]
sklearn_slope = model.coef_[0][1]
print(f'Intercept: {sklearn_intercept}. Slope: {sklearn_slope}')

assert np.abs(sklearn_intercept - normal_intercept) < 0.5
assert np.abs(sklearn_slope - normal_slope) < 0.5
```

## TODO

- Add assertions for model predictions.
- Implement mini-batch SGD.

```python

```