<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/LogisticRegression_Iris.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Associate Flowers With Their Class with Logistic Regression

In this kata, you'll try to predict a flower's class given some of its characteristics.

[Iris](https://archive.ics.uci.edu/ml/datasets/iris) is a well-known multiclass dataset. It contains 3 classes of flowers with 50 examples each. There are a total of 4 features for each flower.

![](images/Iris-versicolor-21_1.jpg)

## Package setup

```python
# The mlkatas package contains various utility functions required by all katas
!pip install mlkatas
```

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
import mlkatas
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 10
%config InlineBackend.figure_format = 'retina'
sns.set()
```

<!-- grade_id: cell-91ddec2f7e596d19, points: 0, grade: false, solution: true -->
```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import load_iris

# YOUR CODE HERE
```

## Step 1: Loading the data

```python
# Load the Iris dataset included with scikit-learn
dataset = load_iris()

# Put data in a pandas DataFrame
df_iris = pd.DataFrame(dataset.data, columns=dataset.feature_names)
# Add target and class to DataFrame
df_iris['target'] = dataset.target
df_iris['class'] = dataset.target_names[dataset.target]
# Show 10 random samples
df_iris.sample(n=10)
```

### Question

Store training input data in a variable named `x` and training targets in a variable named `y`.

<!-- grade_id: cell-68cb077281b12fe7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-c35ec2e0b638d9ec, points: 1, grade: true, solution: false -->
```python
print(f'x: {x.shape}. y: {y.shape}')
print(f'Labels: {y}')

assert x.shape == (150,4)
assert y.shape == (150,)
```

## Step 2: Preparing the data

### Question

Split data and labels into the `x_train`, `x_test`, `y_train` and `y_test` variable using a 20% ratio.

<!-- grade_id: cell-dc46c98f99152d22, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-4a4e8f94ff51de0b, points: 1, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
print(f'x_test: {x_test.shape}. y_test: {y_test.shape}')

assert x_train.shape == (120, 4)
assert y_train.shape == (120,)
assert x_test.shape == (30, 4)
assert y_test.shape == (30,)
```

## Step 3: Training a model

### Question
 
Create a `SGDClassifier` instance and store it into the `model` variable. Fit this model on the training data.

<!-- grade_id: cell-fd837bbf22683090, points: 0, grade: false, solution: true -->
```python
# Train a multiclass classifier on the data

# YOUR CODE HERE
```

## Step 4: Evaluating the model

<!-- grade_id: cell-a93a4cc26df606c8, points: 2, grade: true, solution: false -->
```python
# Compute accuracy on training and test sets
train_acc = model.score(x_train, y_train)
test_acc = model.score(x_test, y_test)

print(f'Training accuracy: {train_acc * 100:.2f}%')
print(f'Test accuracy: {test_acc * 100:.2f}%')
```

## Step 5: Improve results

### Question

Try your best to improve the model's results. You should obtain accuracys around 90%.

<!-- grade_id: cell-b5a8e9870cac09cb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-01a61be5e58aff7e, points: 2, grade: true, solution: false -->
```python
train_acc = model.score(x_train, y_train)
test_acc = model.score(x_test, y_test)

print(f'Training accuracy: {train_acc * 100:.2f}%')
print(f'Test accuracy: {test_acc * 100:.2f}%')

assert train_acc > .85
assert test_acc > .85
```

```python

```