<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/RandomForest_Diabetes.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Predict Diabetes Evolution

In this kata, you'll use several regression models to predict the disease progression one year after.

The Diabetes dataset contains ten baseline variables, age, sex, body mass index, average blood pressure, and six blood serum measurements were obtained for each of n = 442 diabetes patients, as well as the response of interest, a quantitative measure of disease progression one year after baseline.

## Package setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
```

### Question

Import the needed packages.

<!-- grade_id: cell-a3604c4afc82c828, points: 0, grade: false, solution: true -->
```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import load_diabetes

# YOUR CODE HERE
```

## Step 1: Loading the data

```python
dataset = load_diabetes()

# Put data in a pandas DataFrame
df_diab = pd.DataFrame(dataset.data, columns=dataset.feature_names)
# Add target to DataFrame
df_diab['target'] = dataset.target
# Show 10 random samples
df_diab.sample(n=10)
```

## Step 2: Preparing the data

### Question

Split the dataset into training and test sets with a 20% ratio.

<!-- grade_id: cell-4b94ad7a9c319486, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-7135eb33d9bae222, points: 0, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
print(f'x_test: {x_test.shape}. y_test: {y_test.shape}')

assert x_train.shape == (353, 10)
assert y_train.shape == (353,)
assert x_test.shape == (89, 10)
assert y_test.shape == (89,)
```

## Step 3: Training several models

### Question

Create and train a `DecisionTreeRegressor`, a `GradientBoostingRegressor` and a `RandomForestRegressor` on the training data.

Compute their MSE on the training and test data.

<!-- grade_id: cell-f80bcaf317c3e6a7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-0ee3ffaa00916222, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-982fc1c1517d3914, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-7d3465c2a61dab34, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```