> This is a self-correcting activity generated by [nbgrader](https://nbgrader.readthedocs.io). Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# NumPy

## Environment setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 8
%config InlineBackend.figure_format = 'retina'
```

```python
# Import ML packages
import sklearn
print(f'scikit-learn version: {sklearn.__version__}')

from sklearn.datasets import load_sample_images
```

## Part 1: Tensor Basics

### Question

Create a 2D tensor (a matrix) with dimensions (3,4) containing integer values of your choice. Store this tensor in a variable named `x`.

<!-- grade_id: cell-d101d03db303163a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-55a2cf8d88c79c1b, points: 1, grade: true, solution: false -->
```python
print(x)

# Assert dimensions
assert x.ndim == 2
assert x.shape == (3, 4)

# Assert data type 
assert issubclass(x.dtype.type, np.integer)
```

### Question

Update the shape of the previous tensor so that it has dimensions (6,2).

<!-- grade_id: cell-3e614095929cf3c1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-a570a9d5c4254704, points: 1, grade: true, solution: false -->
```python
print(x)

# Assert tensor dimensions
assert x.ndim == 2
assert x.shape == (6, 2)
```

### Question

Change the type of the previous tensor values to `float32`.

<!-- grade_id: cell-221fac9eeab624b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-e522a2bf7881fc07, points: 1, grade: true, solution: false -->
```python
print(x)

# Assert data type
assert issubclass(x.dtype.type, np.floating)
```

## Part 2: Image Management

```python
# Load samples images
images = np.asarray(load_sample_images().images)
print(f'Number of images: {len(images)}. Images tensor: {images.shape}')

first_image = images[0]
# Display first image
plt.imshow(first_image)

# Print details about first image
print(f'First image: {first_image.shape}')
```

### Question

Store in variables respectively named `rgb_values_topleft` and `rgb_values_bottomright` the RGB values of the top-left and bottom-right pixels of the first image.

<!-- grade_id: cell-040daddcf24c1e3e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-7096dc3c71a45431, points: 1, grade: true, solution: false -->
```python
print(f'Top-left pixel: {rgb_values_topleft}')
assert rgb_values_topleft.shape == (3,)

print(f'Bottom-right pixel: {rgb_values_bottomright}')
assert rgb_values_bottomright.shape == (3,)
```

### Question

Reshape the previous `images` tensor into a 2D tensor.

<!-- grade_id: cell-68640d5dfd5fd32d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-dce233b84b3ca086, points: 1, grade: true, solution: false -->
```python
# Assert new tensor dimensions
assert images.shape == (2, 819840)

# Assert RGB values of top-left in first image
assert np.array_equal(rgb_values_topleft, images[0,:3])

# Assert RGB values of bottom-right pixel in first image
assert np.array_equal(rgb_values_bottomright, images[0,819837:])
```