This is a self-correcting activity generated by [nbgrader](https://nbgrader.readthedocs.io).

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Analyse Data With pandas

inspired by an activity in the [Open Machine Learning Course](https://mlcourse.ai/).

## Part 1 : Loading data

```python
# Import base packages
import numpy as np
import pandas as pd
```

### Question

Use Pandas to load the dataset located [here](https://raw.githubusercontent.com/bpesquet/machine-learning-katas/master/datasets/athlete_events.csv), which was extracted from [this Kaggle competition](https://www.kaggle.com/heesoo37/120-years-of-olympic-history-athletes-and-results/). Store it into a DataFrame named `df_olympics`.

<!-- grade_id: cell-6930168ca952d16e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a4b2d9d694a4d11b, points: 1, grade: true, solution: false -->
```python
print(f'df_olympics: {df_olympics.shape}')

assert isinstance(df_olympics, pd.DataFrame)
assert df_olympics.shape == (271116, 15)
```

## Part 2: Exploring data

### Question

Print general information about the DataFrame.

<!-- grade_id: cell-cc980d2fe4eb912e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question

Print and observe the first 10 records of the dataset.

<!-- grade_id: cell-c2348e0b371280d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question

Print and observe 5 random records of the dataset.

<!-- grade_id: cell-ef82a27734491ed4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```