> This is a self-correcting activity generated by [nbgrader](https://nbgrader.readthedocs.io). Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Classify handwritten digits

In this activity, you'll try several classifiers on the [UCI handwritten digits dataset](https://archive.ics.uci.edu/ml/datasets/Optical+Recognition+of+Handwritten+Digits), either separately or into an ensemble.

![UCI digits](images/uci_digits.png)

## Environment setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 8
%config InlineBackend.figure_format = 'retina'
sns.set()
```

```python
# Import ML packages
import sklearn
print(f'scikit-learn version: {sklearn.__version__}')

from sklearn.datasets import load_digits
from sklearn.linear_model import SGDClassifier
from sklearn.ensemble import RandomForestClassifier
from sklearn.neural_network import MLPClassifier
from sklearn.ensemble import VotingClassifier
from sklearn.model_selection import train_test_split
```

## Step 1: Loading and preparing the data

```python
# Load the MNIST digits dataset
digits = load_digits()

# To apply a classifier on this data, we need to flatten the image, to
# turn the data in a (samples, feature) matrix:
n_samples = len(digits.images)
data = digits.images.reshape((n_samples, -1))

print(f'data: {data.shape}. targets: {digits.target.shape}')
```

### Question

Split the data into training, validation and test sets.

<!-- grade_id: cell-fb6a9c746b43c1af, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-60dcf8104737b392, points: 2, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. x_val: {x_val.shape}. x_test: {x_test.shape}')
print(f'y_train: {y_train.shape}. y_val: {y_val.shape}. y_test: {y_test.shape}')

assert x_train.shape == (1237, 64)
assert x_val.shape == (200, 64)
assert x_test.shape == (360, 64)
assert y_train.shape == (1237,)
assert y_val.shape == (200,)
assert y_test.shape == (360,)
```

## Step 2: Train several models

### Question

Create and train various models, such as a linear classifier, a multilayer perceptron, a random forest...

<!-- grade_id: cell-9cf3e82d97c98c75, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Display the default score for each model.

<!-- grade_id: cell-4f60774a06371248, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Create a `VotingClassifier` including all your models. Fit it on the training data.

<!-- grade_id: cell-ccc315abafd8246f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Show the `VotingClassifier` score and compare to each model's individual score.

<!-- grade_id: cell-560a77d9f9b57409, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Show the score for a soft voting classifier.

<!-- grade_id: cell-ff3fc8fba60b9acd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

### Question

Compute the `VotingClassifier` score on the test data. Compare it to each model's individual score.

<!-- grade_id: cell-7a9a6c423c33ba0d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```